%lang starknet

from starkware.cairo.common.serialize import serialize_word

## I AM NOT DONE

# Create a function that accepts a parameter and logs it
func logValue(y : felt):      

   # Create a hint that logs to console value of y


   ## Solution
   # %{ 
   #    print(f"passed value: {ids.y}")      
   # %}   

   return ()   
end



