%lang starknet

from starkware.cairo.common.serialize import serialize_word

## I AM NOT DONE

# Create a function that accepts a parameter and logs it
func logValue(y : felt):      

   # Create a function to add one to a value  
   %{ 
        print(f"passed value: {ids.y}")
      # print("\n\n666666\n\n")
   %}

   let u = 9

   return ()   
end



