%lang starknet

## I AM NOT DONE

from starkware.cairo.common.serialize import serialize_word

# Create a function to add one to the value passed
func addOne(y : felt) -> (bit : felt):

   ## uncomment this

   # Solution ex1
   ################################################################################################
   return (y + 1)
   ################################################################################################
end


